# Module 4 Assignment

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `# YOUR CODE HERE`. Do not write your answer anywhere else other than where it says `# YOUR CODE HERE`. Anything you write elsewhere will be removed or overwritten by the autograder.
2. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and restart the kernel and run all cells (Restart & Run all).
3. Do not change the title (i.e. file name) of this notebook.
4. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).
5. All work must be your own, if you do use any code from another source (such as a course notebook or a website) you need to properly cite the source.

-----

```python
%matplotlib inline

import pandas as pd
import numpy as np

import matplotlib.pyplot as plt
import seaborn as sns

from sklearn.model_selection import train_test_split

from nose.tools import assert_equal, assert_almost_equal, assert_true, assert_is_instance
```

-----

## Loading Breast Cancer Data

In this assignment, we will work with a breast cancer data set to make predictive models. Before we build a model, we first load the data into the assignment notebook, and randomly sample several rows. The second Code cell splits the DataFrame into a training and testing data set, respectively, before creating the features and labels to use for our classification task.

-----

```python
# Load data
df = pd.read_csv('breast-cancer-wisconsin.data')
df.sample(5)
```

Output:
```
id  clump thickness  uniformity cell size  uniformity cell shape  \
369  1022257                2                     1                      1   
155  1196915               10                     7                      7   
330   822829                7                     6                      4   
161  1198128               10                     8                     10   
545  1321942                5                     1                      1   

     marginal adhesion  epithelial cell size  bare nuclei  bland chromatin  \
369                  1                     2            1                1   
155                  4                     5           10                5   
330                  8                    10           10                9   
161                 10                     6            1                3   
545                  1                     2            1                3   

     normal nucleoli  mitoses  class  
369                1        1    
```

```python
# Create features and labels

y = df['class']
X = df.drop('class', axis=1)

# Split data into training and testing data sets
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.3, random_state=0)
```

```python
# Load data
df = pd.read_csv('./breast-cancer-wisconsin.data')
df.sample(5)
```

```python
# Create features and labels

y = df['class']
X = df.drop('class', axis=1)

# Split data into training and testing data sets
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.3, random_state=0)
```

-----

## Problem 1: Create a Pipeline for Random Forest

For this problem, you must complete the `rf_pipeline` function to create a new pipeline that contains random forest classifier. This classifier should be created by using the `n_estimators` and `rs` parameters.

-----

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import RandomForestClassifier
from sklearn.pipeline import Pipeline

def rf_pipeline(n_estimators, rs):
    """"
    Create a scikit learn pipeline that wraps a random forest classifier
    
    Inputs
    ---------
    n_estimators: the number of estimators used by the random forest
    rs: the random state of the random forest, used for reproducibility
    
    Returns
    -------
    pipeline: an sklearn pipeline object
    """
    
    ### YOUR CODE HERE
    
    return pipeline
```

```python
rf = RandomForestClassifier(n_estimators=100, random_state=2)
```

<!-- grade_id: p1_test, points: 1, grade: true, solution: false -->
```python
# Call function to create a pipeline
pl = rf_pipeline(100, 2)

# Test the new pipeline
assert_equal(str(type(pl)), "<class 'sklearn.pipeline.Pipeline'>")
assert_equal(pl.steps[0][1].n_estimators, 100)
assert_equal(pl.steps[0][1].random_state, 2)
```

-----

## Problem 2: Compute the Average Cross Validation Score

For this problem, you will finish the `rf_avg_crossval` function to compute the average (or mean) cross validation score for a random forest classifier. To perform the cross-validation, you should use the `StratifiedKFold` object within the scikit learn library. Specifically, you must complete the following tasks:

1. Create the `StratifiedKFold` object, using the `n_splits` and `rs` parameters.
2. Create a random forest pipeline, by using your solution to problem one above, and including the `n_estimators` and `rs` parameters.
3. Compute the cross-validation scores for this pipeline, the training data, and the `StratifiedKFold` object.
4. Compute the average cross validation score

-----

<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import StratifiedKFold, cross_val_score

def rf_avg_crossval(rfp, X_train, y_train, 
                    n_estimators, n_splits, rs):
    """"
    Inputs
    ---------
    rfp: The rf_pipeline function you created in Problem 1
    X_train: the training data
    y_train: the training data labels
    n_estimators: the number of estimators used by the random forest
    rs: the random state of the random forest and of the cross validation,
        used for reproducibility
    
    Returns
    -------
    mean_score: the average cross validation score of the random forest 
    """
    
    ### YOUR CODE HERE
    
    return mean_score
```

```python
skf = StratifiedKFold(n_splits=5, random_state=2, shuffle=True)
```

```python
rf = RandomForestClassifier(n_estimators=10, random_state=2)
```

```python
rf.fit(X_train,y_train)
```

Output:
```
RandomForestClassifier(n_estimators=10, random_state=2)
```

```python
cross_val_score(rf,X_test,y_test)
```

Output:
```
array([0.95121951, 0.97560976, 0.95121951, 0.92682927, 0.97560976])
```

```python
cross_val_score(rf,X_test,y_test).mean()
```

Output:
```
0.9560975609756097
```

<!-- grade_id: p2_test, points: 1, grade: true, solution: false -->
```python
# Compute mean scores
mean_score = rf_avg_crossval(rf_pipeline, X_train, y_train, 10, 5, 2)

# Test mean scores
assert_almost_equal(mean_score, 0.96236, places=4)
```

-----

## Problem 3: Random Grid Search

For this problem, you will finish writing the `random_crossval` function, shown below, to complete the following tasks:

1. Create a dictionary to map the `'rf__n_estimators'` key to the `n_estimators_array` value.
1. Create a randomized search cross-validation estimator by using the scikit learn library. Be sure to pass or assign the following parameters when creating this estimator:
  - `pl` random forest pipeline.
  - The dictionary created in step one above assigned to the `param_distributions` parameter.
  - `n_samples` parameter assigned to the `n_iter` parameter.
  - `rs` assigned to the `random_state` parameter.

1. Fit this estimator to the training data.
1. Extract the best values from the randomized grid search as a dictionary.
1. Return the dictionary containing the best parameters.

-----

```python
from sklearn.ensemble import RandomForestClassifier
rf = RandomForestClassifier(random_state=2)
```

<!-- grade_id: p3_answer, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import RandomizedSearchCV

def random_crossval(pl, X_train, y_train, n_estimators_array, n_samples, rs):
    """"
    Uses a randomized grid search to determine the best parameters of a 
    specific the specified pipeline on a given test data set.
    
    Inputsb
    ---------
    pl: random forest pipeline
    X_train: the training data
    y_train: the trainig data lavels
    n_estimators_array: a numpy array containing possible values for n_estimators
    n_samples: the number of times the random cv will sample from n_estimators_array
    rs: the random state of the random forest and of the cross validation,
        used for reproducibility
    
    Returns
    -------
    best_params: a list of the best parameters from the randomized cross validation
    """
    
    ### YOUR CODE HERE
    
    return best_params
```

```python
parameters = {'n_estimators':np.arange(1,100,1)
             }
```

```python
randm = RandomizedSearchCV(estimator=rf, param_distributions = parameters, cv = 5, n_iter = 10, 
                           n_jobs=-1, scoring=None)
```

```python
randm.fit(X_train,y_train)
```

Output:
```
RandomizedSearchCV(cv=5, estimator=RandomForestClassifier(random_state=2),
                   n_jobs=-1,
                   param_distributions={'n_estimators': array([ 1,  2,  3,  4,  5,  6,  7,  8,  9, 10, 11, 12, 13, 14, 15, 16, 17,
       18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30, 31, 32, 33, 34,
       35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49, 50, 51,
       52, 53, 54, 55, 56, 57, 58, 59, 60, 61, 62, 63, 64, 65, 66, 67, 68,
       69, 70, 71, 72, 73, 74, 75, 76, 77, 78, 79, 80, 81, 82, 83, 84, 85,
       86, 87, 88, 89, 90, 91, 92, 93, 94, 95, 96, 97, 98, 99])})
```

```python
randm.best_params_
```

Output:
```
{'n_estimators': 32}
```

<!-- grade_id: p3_test, points: 1, grade: true, solution: false -->
```python
# Perform Grid Search to get best parameters
params_list = random_crossval(pl, X_train, y_train, 
                              np.arange(1,100,1), 10, 2)

# Test grid search results
assert_equal(params_list["rf__n_estimators"], 31)
```

-----

## Loading Car Price Data

For the rest of this assignment, we will work with a car price data set to make regressive models between car specifications and car price. Before we build a model, we first load the data into the assignment notebook, and randomly sample several rows. The second Code cell selects only numeric columns to simplify the analysis, and displays the first few instances.

-----

```python
df = pd.read_csv('imports-85.data')
df.sample(5)
```

Output:
```
symboling  normalized-losses  make  fuel-type aspiration  num-of-doors  \
153          0                 77    19          1        std           4.0   
198         -2                103    21          1      turbo           4.0   
82           3                122    11          1      turbo           2.0   
148          0                 85    18          1        std           4.0   
49           0                122     7          1        std           2.0   

     body-style drive-wheels 'engine-location'  wheel-base  ...  engine-size  \
153           1          fwd             front        95.7  ...           92   
198           2          rwd             front       104.3  ...          130   
82            3          fwd             front        95.9  ...          156   
148           1          4wd             front        96.9  ...          108   
49            2          rwd             front       102.0  ...          326   

     fuel-system  bore  stroke compression-ratio'
```

```python
df_simple = df[df.columns[df.dtypes!=object]]
df_simple.head()
```

Output:
```
symboling  normalized-losses  make  fuel-type  num-of-doors  body-style  \
0          3                122     0          1           2.0           4   
1          3                122     0          1           2.0           4   
2          1                122     0          1           2.0           3   
3          2                164     1          1           4.0           2   
4          2                164     1          1           4.0           2   

   wheel-base  length  width  height  curb-weight  engine-size  \
0        88.6   168.8   64.1    48.8         2548          130   
1        88.6   168.8   64.1    48.8         2548          130   
2        94.5   171.2   65.5    52.4         2823          152   
3        99.8   176.6   66.2    54.3         2337          109   
4        99.4   176.6   66.4    54.3         2824          136   

   compression-ratio'  city-mpg  highway    price  
0                 9.0        21       27  13495.0  
1                 9.0        21  
```

```python
df = pd.read_csv('./imports-85.data')
df.sample(5)
```

```python
df_simple = df[df.columns[df.dtypes!=object]]
df_simple.head()
```

-----

With the numeric data, we can now create our independent and dependent variables so that we can 
build a model that can predict the price of the car. After this is done, we split these data into a training and testing set.

-----

```python
X = df_simple.drop("price", axis=1)
y = np.ravel(df_simple["price"])
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.3, random_state=0)
```

```python
X_train, y_train
```

Output:
```
(     symboling  normalized-losses  make  fuel-type  num-of-doors  body-style  \
 40           0                 85     5          1           4.0           2   
 60           0                115     8          1           4.0           2   
 56           3                150     8          1           2.0           3   
 101          0                128    12          1           4.0           2   
 86           1                125    11          1           4.0           2   
 ..         ...                ...   ...        ...           ...         ...   
 67          -1                 93     9          0           4.0           2   
 192          0                122    20          0           4.0           2   
 117          0                161    13          1           4.0           2   
 47           0                145     7          1           4.0           2   
 172          2                134    19          1           2.0           4   
 
      wheel-base  length  
```

-----

## Problem 4:  Ridge Regression

For this problem, you must finish writing the `ridge_r` function by completing the following tasks:
1. Create a [`Ridge Regression`](http://scikit-learn.org/stable/modules/generated/sklearn.linear_model.Ridge.html) model by using the scikit learn library. Be sure to
  - Set the `alpha` hyperparameter for the `RidgeRegression` model to the `alpha` function parameter.
  - Set the `random_state` parameter for the `RidgeRegression` model to the `rs` function parameter.
  
1. Fit the `RidgeRegression` model by using the training features and labels.
1. Compute the $R^2$ coefficient using the testing features and labels.
1. Return the $R^2$ coefficient.

-----

```python
from sklearn.linear_model import Ridge
```

```python
ridge_model = Ridge(alpha=1, random_state=0)
```

```python
ridge_model.fit(X_train, y_train)
```

Output:
```
Ridge(alpha=1, random_state=0)
```

```python
y_pred = ridge_model.predict(X_test)
```

```python
from sklearn.metrics import r2_score
```

```python
r2_score(y_test,y_pred)
```

Output:
```
0.6781910609377904
```

```python

```

```python
ridge_model = Ridge(alpha=250, random_state=0)
```

```python
ridge_model.fit(X_train, y_train)
```

Output:
```
Ridge(alpha=250, random_state=0)
```

```python
y_pred = ridge_model.predict(X_test)
```

```python
from sklearn.metrics import r2_score
```

```python
r2_score(y_test,y_pred)
```

Output:
```
0.733134627419015
```

<!-- grade_id: p4_answer, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import Ridge

def ridge_r(X_train, y_train, X_test, y_test, alpha=1, rs=0):
    """"
    Fit a ridge regression model to the training data and 
    compute and return the R2 score for the test data.
    
    Inputs
    ---------
    X_train: Pandas DataFrame containing features for training dataset
    y_train: Pandas DataFrame containing labels for training dataset
    X_test: Pandas DataFrame containing features for testing dataset
    y_test: Pandas DataFrame containing labels for testing dataset
    alpha: a regularization parameter
    rs: random state for random number generator
    
    Returns
    -------
    score: the coefficient of determination as a floating point value
    """
    
    ### YOUR CODE HERE
```

<!-- grade_id: p4_test, points: 1, grade: true, solution: false -->
```python
import helper

# Compute scores as a function of alpha
alphas = np.arange(0.1, 250, 1)
scores = []
for alpha in alphas:
    score = ridge_r(X_train, y_train, X_test, y_test, alpha=alpha)
    scores.append(score)

# Compare scores to precomputed values
for ans, sol in zip(scores, helper.rr_sol):
    assert_almost_equal(ans, sol, places=2)
```

-----

If you plot your ridge regression scores against alpha they should look similar to this plot: 

![ridgeR2.png](ridgeR2.png)

Notice how the r2 score increases as alpha increases. 

*If you are failing the previous assert statement, try running the cell below to help you debug the `ridge_r` function.*

-----

```python
# Plot the score as a function of alpha
sns.set()
fig, ax = plt.subplots()
ax.plot(alphas, scores)
ax.set(title='Ridge Regression Score',
       xlabel='Alpha',
       ylabel='Score')

sns.despine()
```

**&copy; 2017: Robert J. Brunner at the University of Illinois.**

This notebook is released under the [Creative Commons license CC BY-NC-SA 4.0][ll]. Any reproduction, adaptation, distribution, dissemination or making available of this notebook for commercial use is not allowed unless authorized in writing by the copyright holder.

[ll]: https://creativecommons.org/licenses/by-nc-sa/4.0/legalcode 